```python
version = "REPLACE_PACKAGE_VERSION"
```

# Experiment Design and Analysis
## School of Information, University of Michigan

## Week 4: 
- 1. Threats to Validity
- 2. Instrumental Variables

## Assignment Overview
### The objective of this assignment is to:

- Apply theory of experiment design and knowledge of analysis techniques to real experiment data.


### The total score of this assignment will be 25 points


### Resources:
- StatsModels
    - We recommend using a python library called [StatsModels](https://www.statsmodels.org/stable/index.html) for data analysis


- Dataset used in this assignment: Kiva Crowdsourcing Team data [view source files](https://www.openicpsr.org/openicpsr/project/100358/version/V2/view)
    - Source for dataset: [Chen, Y., et al. Recommending teams promotes prosocial lending in online microfinance (2016).](https://www.pnas.org/content/113/52/14944)

```python
import pandas as pd
import numpy as np
import statsmodels.api as sm
from linearmodels.iv import IV2SLS #you may get a deprecation warning for this library -- this is fine.

data = pd.read_csv('assets/assignment4_data.csv') #Data for this assignment
```

```python
#uncomment the below line to view readme files for this dataset (includes explanation of variable names)
# !cat assets/assignment4_data_readme.md

# #uncomment the below line to view snippet of csv file
# data.head()
```

## Part A (15 points)

We want to assess the effectiveness of joining a team on Kiva -- specifically, what impact joining a team has on donations. Using the variable indicating whether users have joined a team (```join```) and the differences in donations made over a certain period (```amt_diff_1d```, ```amt_diff_7d```, ```amt_diff_30d```), we can find if joining a team has impact on donations. However, variables that determine whether subjects join a team may also affect the amount they donate since we only inform subjects about joining a team.

In this case, our instrumental variable is whether the subjects were sent an e-mail to inform about the team functionality on Kiva. 

***Before you go on, recall from lecture the requirements an instrumental variable must satisfy - you will be investigating these requirements in this notebook.***

1. To get started, we need to create the instrumental variable. Add a new column named ```email``` in the dataframe. The value for ```email``` should be ```1``` if users received an e-mail as part of their treatment group (```treatment_id``` != 1), and ```0``` if they did not (```treatment_id``` = 1). (2 points)

<!-- grade_id: add_email_df, points: 0, grade: false, solution: true -->
```python
data.rename(columns={'join': 'join_any'}, inplace=True) #since join is also the name of a pandas method, we rename the column to avoid confusion
# YOUR CODE HERE
data["email"] = (data["treatment_id"]!=1).astype(int)
data
```

Output:
```
shuffled_lender_id  treatment_id  join_any  join_rec  opened  \
0                       0             7         0         0       1   
1                       1             3         0         0       1   
2                       2             4         0         0       1   
3                       3             2         0         0       0   
4                       4             5         0         0       0   
...                   ...           ...       ...       ...     ...   
64795               64795             4         0         0       1   
64796               64796             3         0         0       0   
64797               64797             1         0         0       0   
64798               64798             7         0         0       0   
64799               64799             1         0         0       0   

       amt_diff_1d  amt_diff_7d  amt_diff_30d  email  
0              0.0          0.0         -50.0      1  
1              0.0          0.0           0.
```

```python

```

<!-- grade_id: no_email_nans, points: 1, grade: true, solution: false -->
```python
assert pd.notnull(data['email'].all()), "email column must contain either 0 or 1"
```

<!-- grade_id: all_treat_groups, points: 1, grade: true, solution: false -->
```python
assert data.loc[data['treatment_id'] != 1,'email'].all() == 1, "all treatments except treatment 1 received an email"
assert data.loc[data['treatment_id'] == 1,'email'].all() == 0, "all treatments except treatment 1 received an email"
```

2. Next, we will create a constant, equal to 1. Add a column in the dataframe called ```const```. (1 point)

<!-- grade_id: add_const_df, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
data["const"]=1
```

<!-- grade_id: const_df_added, points: 1, grade: true, solution: false -->
```python
assert data['const'].all() == 1, "the constant value should be 1"
```

In lecture, the 2-stage least squares model was used. Now, let’s follow the steps indicated in lecture to create this model to measure the effect described above. First, we need to estimate the effect of e-mailing users to join a team on whether they join a team.


3. Using statsmodels, create an ordinary least squares regression model that does this. Fit the model and store it in the variable: ```model_fs```. Using the predict method from ```model_fs```, store the predicted values in a new column in your dataframe called ```predicted_join```. Recall from lecture that since we have created this new variable, we can estimate the effect of joining a team on lending amounts without worrying about the effect of potential unobserved or missing variables. (4 points)

Note: ensure your model has a constant.

```python
data
```

Output:
```
shuffled_lender_id  treatment_id  join_any  join_rec  opened  \
0                       0             7         0         0       1   
1                       1             3         0         0       1   
2                       2             4         0         0       1   
3                       3             2         0         0       0   
4                       4             5         0         0       0   
...                   ...           ...       ...       ...     ...   
64795               64795             4         0         0       1   
64796               64796             3         0         0       0   
64797               64797             1         0         0       0   
64798               64798             7         0         0       0   
64799               64799             1         0         0       0   

       amt_diff_1d  amt_diff_7d  amt_diff_30d  email  const  
0              0.0          0.0         -50.0      1      1  
1              0.0          0.
```

<!-- grade_id: ols_model_predicted_join, points: 0, grade: false, solution: true -->
```python
def email_join_ols(provided_data):
    # YOUR CODE HERE
    X = provided_data[["email", "const"]].values
    y = provided_data["join_any"]
    res = sm.Logit(y,X).fit()
    provided_data["predicted_join"] = res.predict(X)
    return provided_data
```

Your function should return a dataframe with the correct values and columns. Check that it does:

```python
email_join_ols(data).head() #you may get a deprecation warning here -- that is ok.
```

Output:
```
Optimization terminated successfully.
         Current function value: 0.051557
         Iterations 9
   shuffled_lender_id  treatment_id  join_any  join_rec  opened  amt_diff_1d  \
0                   0             7         0         0       1          0.0   
1                   1             3         0         0       1          0.0   
2                   2             4         0         0       1          0.0   
3                   3             2         0         0       0          0.0   
4                   4             5         0         0       0          0.0   

   amt_diff_7d  amt_diff_30d  email  const  predicted_join  
0          0.0         -50.0      1      1        0.009801  
1          0.0           0.0      1      1        0.009801  
2          0.0           0.0      1      1        0.009801  
3          0.0           0.0      1      1        0.009801  
4          0.0           0.0      1      1        0.009801
```

<!-- grade_id: predicted_join_col, points: 1, grade: true, solution: false -->
```python
assert 'predicted_join' in data, "checking there is a column named predicted_join in data"
```

<!-- grade_id: correct_pred_join_treat, points: 3, grade: true, solution: false -->
```python
"""checking the correct predicted_join values are present"""
# Hidden tests
```

Output:
```
'checking the correct predicted_join values are present'
```

Now that we have the predicted values of whether a subject would be expected to join a team based on if they were e-mailed, we can move to the second stage.

4. In this stage, we will run the estimation of the effect of joining a team on the amount a subject lends. However, instead of using the ```join``` variable, we will use our new ```predicted_join``` variable. Using statsmodels again, and ensuring your model has a constant, create three ordinary least squares regression models which estimate the effect of the prediction of users joining a team on the following:

a. ```amt_diff_1d```, storing the fitted model in ```model_1d``` (3 points)

<!-- grade_id: amt_1d_ols, points: 0, grade: false, solution: true -->
```python
def pred_join_amt_1d(provided_data):
    # YOUR CODE HERE
    X = provided_data[["predicted_join", "const"]].values
    y = provided_data["amt_diff_1d"]
    res = sm.OLS(y,X).fit()
    return res
```

Your function should return a summary view of your results. Check that it does:

```python
print(pred_join_amt_1d(data).summary()) #we've wrapped this in a print statement to preserve the original statsmodels layout. You may get a deprecation warning here -- that is ok.
```

Output:
```
OLS Regression Results                            
==============================================================================
Dep. Variable:            amt_diff_1d   R-squared:                       0.001
Model:                            OLS   Adj. R-squared:                  0.001
Method:                 Least Squares   F-statistic:                     56.75
Date:                Wed, 23 Mar 2022   Prob (F-statistic):           5.01e-14
Time:                        11:52:34   Log-Likelihood:            -2.8014e+05
No. Observations:               64800   AIC:                         5.603e+05
Df Residuals:                   64798   BIC:                         5.603e+05
Df Model:                           1                                         
Covariance Type:            nonrobust                                         
==============================================================================
                 coef    std err          t      P>|t|      [0.025      0.975]
-
```

<!-- grade_id: correct_amt_1d_ols, points: 3, grade: true, solution: false -->
```python
"""checking your t-value is correct"""
# Hidden tests
```

Output:
```
'checking your t-value is correct'
```

b. ```amt_diff_7d```, storing the fitted model in ```model_7d``` (3 points)

<!-- grade_id: amt_7d_ols, points: 0, grade: false, solution: true -->
```python
def pred_join_amt_7d(provided_data):
    # YOUR CODE HERE
    X = provided_data[["predicted_join", "const"]].values
    y = provided_data["amt_diff_7d"]
    res = sm.OLS(y,X).fit()
    return res
```

Your function should return a summary view of your results. Check that it does:

```python
print(pred_join_amt_7d(data).summary())
```

Output:
```
OLS Regression Results                            
==============================================================================
Dep. Variable:            amt_diff_7d   R-squared:                       0.000
Model:                            OLS   Adj. R-squared:                  0.000
Method:                 Least Squares   F-statistic:                     9.979
Date:                Wed, 23 Mar 2022   Prob (F-statistic):            0.00158
Time:                        11:53:15   Log-Likelihood:            -3.5400e+05
No. Observations:               64800   AIC:                         7.080e+05
Df Residuals:                   64798   BIC:                         7.080e+05
Df Model:                           1                                         
Covariance Type:            nonrobust                                         
==============================================================================
                 coef    std err          t      P>|t|      [0.025      0.975]
-
```

<!-- grade_id: correct_amt_7d_ols, points: 3, grade: true, solution: false -->
```python
"""checking your t-value is correct"""
# Hidden tests
```

Output:
```
'checking your t-value is correct'
```

c. ```amt_diff_30d```, storing the fitted model in ```model_30d``` (2 points)

<!-- grade_id: amt_30d_ols, points: 0, grade: false, solution: true -->
```python
def pred_join_amt_30d(provided_data):
    # YOUR CODE HERE
    X = provided_data[["predicted_join", "const"]].values
    y = provided_data["amt_diff_30d"]
    res = sm.OLS(y,X).fit()
    return res
```

Your function should return a summary view of your results. Check that it does:

```python
print(pred_join_amt_30d(data).summary())
```

Output:
```
OLS Regression Results                            
==============================================================================
Dep. Variable:           amt_diff_30d   R-squared:                       0.000
Model:                            OLS   Adj. R-squared:                  0.000
Method:                 Least Squares   F-statistic:                     2.112
Date:                Wed, 23 Mar 2022   Prob (F-statistic):              0.146
Time:                        11:53:53   Log-Likelihood:            -3.8856e+05
No. Observations:               64800   AIC:                         7.771e+05
Df Residuals:                   64798   BIC:                         7.771e+05
Df Model:                           1                                         
Covariance Type:            nonrobust                                         
==============================================================================
                 coef    std err          t      P>|t|      [0.025      0.975]
-
```

<!-- grade_id: correct_amt_30d_ols, points: 2, grade: true, solution: false -->
```python
"""checking your t-value is correct"""
# Hidden tests
```

Output:
```
'checking your t-value is correct'
```

## Part B (10 points)

Now we have estimated the effect of joining a team on lending using instrumental variables! However, there is a more direct way to complete these two stages.

Using the IV2SLS (Instrumental Variables 2-Stage Least Squares) function ([Documentation](https://bashtage.github.io/linearmodels/iv/iv/linearmodels.iv.model.IV2SLS.html)) in the linearmodels library will achieve everything we did above faster -- and it will more correctly estimate the standard errors.

1. First, let's make things simpler for ourselves. For this analysis, we will only need a dataframe with the following columns: ```email```, ```join_any```, ```amt_diff_1d```, ```amt_diff_7d```, ```amt_diff_30d```, and the ```const``` column you created above. (As in our other models, we need a constant included in the models we will be creating with IV2SLS.) (1 point)

<!-- grade_id: iv_dataframe, points: 0, grade: false, solution: true -->
```python
iv_dataframe = pd.DataFrame()
# YOUR CODE HERE
iv_dataframe = data[["email", "join_any", "amt_diff_1d", "amt_diff_7d", "amt_diff_30d", "const"]].copy()
```

```iv_dataframe``` should yield a dataframe with the correct calculated, given, and new column and row values. Check that it does:

```python
iv_dataframe.head()
```

Output:
```
email  amt_diff_1d  amt_diff_7d  amt_diff_30d  const
0      1          0.0          0.0         -50.0      1
1      1          0.0          0.0           0.0      1
2      1          0.0          0.0           0.0      1
3      1          0.0          0.0           0.0      1
4      1          0.0          0.0           0.0      1
```

<!-- grade_id: correct_iv_dataframe, points: 1, grade: true, solution: false -->
```python
"""checking your dataframe columns are all present"""
assert 'email' and 'join_any'and 'amt_diff_1d' and 'amt_diff_7d' and 'amt_diff_30d'and 'const' in iv_dataframe
```

After looking over the documentation of the IV2SLS function, create and fit three models (with the three lending measurement periods used in number 3) that estimate the effect of joining a team on lending amounts considering the instrument of emailing subjects about joining a team.

According to the [linearmodels IV2SLS documentation](https://bashtage.github.io/linearmodels/iv/iv/linearmodels.iv.model.IV2SLS.html), you will need to provide a given set of parameters to the function in order to create the model. The dependent variables and instruments are straightforward, but what are exogenous and endogenous regressors?

An exogenous regressor does not co-vary with the model’s random error while an endogenous regressor does. In our model’s case, we know the ```join_any``` variable co-varies with the random error while ```const``` cannot (since it is a constant!).

2. Create and fit the model for the first lending measurement period, the period referred to in ```amt_diff_1d``` (3 points)

<!-- grade_id: amt_1d_iv, points: 0, grade: false, solution: true -->
```python
def iv_model_1d(provided_data):
    
    """ Take some time to think about what exactly you're modeling here, then read the linearmodels documentation.
    What is the instrument? What is the dependent variable, what are the endogenous and exogenous regressors?
    Tip: the covariance should be unadjusted in this model, (and your following models)
    """
    # iv_result_1d = your code here
    # YOUR CODE HERE
    endog = provided_data[["join_any"]]
    exog = provided_data[["const"]]
    d = provided_data["amt_diff_1d"]
    iv_result_1d = IV2SLS(d, exog, endog, provided_data["email"]).fit()
    return iv_result_1d
```

Your function should return a summary view of your results. Check that it does:

```python
iv_model_1d(iv_dataframe)
```

Output:
```
IV-2SLS Estimation Summary                          
==============================================================================
Dep. Variable:            amt_diff_1d   R-squared:                     -2.3237
Estimator:                    IV-2SLS   Adj. R-squared:                -2.3237
No. Observations:               64800   F-statistic:                    13.763
Date:                Wed, Mar 23 2022   P-value (F-stat)                0.0002
Time:                        12:06:26   Distribution:                  chi2(1)
Cov. Estimator:                robust                                         
                                                                              
                             Parameter Estimates                              
==============================================================================
            Parameter  Std. Err.     T-stat    P-value    Lower CI    Upper CI
------------------------------------------------------------------------------
```

<!-- grade_id: correct_amt_1d_iv, points: 3, grade: true, solution: false -->
```python
"""checking your 1d model has an unadjusted covariance and your p-value is correct"""
# Hidden tests
```

3. Create and fit the model for the first lending measurement period, the period referred to in ```amt_diff_7d``` (3 points)

<!-- grade_id: amt_7d_iv, points: 0, grade: false, solution: true -->
```python
def iv_model_7d(provided_data):
    # YOUR CODE HERE
    endog = provided_data[["join_any"]]
    exog = provided_data[["const"]]
    d = provided_data["amt_diff_7d"]
    iv_result_7d = IV2SLS(d, exog, endog, provided_data["email"]).fit()
    return iv_result_7d
```

Your function should return a summary view of your results. Check that it does:

```python
iv_model_7d(iv_dataframe)
```

Output:
```
IV-2SLS Estimation Summary                          
==============================================================================
Dep. Variable:            amt_diff_7d   R-squared:                     -0.4152
Estimator:                    IV-2SLS   Adj. R-squared:                -0.4153
No. Observations:               64800   F-statistic:                    4.2358
Date:                Wed, Mar 23 2022   P-value (F-stat)                0.0396
Time:                        12:11:28   Distribution:                  chi2(1)
Cov. Estimator:                robust                                         
                                                                              
                             Parameter Estimates                              
==============================================================================
            Parameter  Std. Err.     T-stat    P-value    Lower CI    Upper CI
------------------------------------------------------------------------------
```

<!-- grade_id: correct_amt_7d_iv, points: 3, grade: true, solution: false -->
```python
"""checking your 7d model has an unadjusted covariance and your p-value is correct"""
# Hidden tests
```

4. Create and fit the model for the first lending measurement period, the period referred to in ```amt_diff_30d``` (3 points)

<!-- grade_id: amt_30d_iv, points: 0, grade: false, solution: true -->
```python
def iv_model_30d(provided_data):
    # YOUR CODE HERE
    endog = provided_data[["join_any"]]
    exog = provided_data[["const"]]
    d = provided_data["amt_diff_30d"]
    iv_result_30d = IV2SLS(d, exog, endog, provided_data["email"]).fit()
    return iv_result_30d
```

Your function should return a summary view of your results. Check that it does:

```python
iv_model_30d(iv_dataframe)
```

Output:
```
IV-2SLS Estimation Summary                          
==============================================================================
Dep. Variable:           amt_diff_30d   R-squared:                     -0.0806
Estimator:                    IV-2SLS   Adj. R-squared:                -0.0807
No. Observations:               64800   F-statistic:                    1.5562
Date:                Wed, Mar 23 2022   P-value (F-stat)                0.2122
Time:                        12:11:48   Distribution:                  chi2(1)
Cov. Estimator:                robust                                         
                                                                              
                             Parameter Estimates                              
==============================================================================
            Parameter  Std. Err.     T-stat    P-value    Lower CI    Upper CI
------------------------------------------------------------------------------
```

<!-- grade_id: correct_amt_30d_iv, points: 3, grade: true, solution: false -->
```python
"""checking your 30d model has an unadjusted covariance and your p-value is correct"""
# Hidden tests
```

Output:
```
'checking your 30d model has an unadjusted covariance and your p-value is correct'
```

```python

```